---
title        : Chapter 4
description  : Description 

--- type:VideoExercise lang:r xp: skills: key:a05bc8eece
## Exercise 1 
Content

*** =video_link

--- type:NormalExercise lang:r xp: skills: key:a9fea4cad7
## Exercise 2 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp: skills: key:4eba185201
## Exercise 3 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:MultipleChoiceExercise lang:r xp: skills: key:f9bf75f28c
## Exercise 4 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```

--- type:VideoExercise lang:r xp: skills: key:1b42625ac2
## Exercise 5 
Content

*** =video_link

--- type:NormalExercise lang:r xp: skills: key:a669c2afce
## Exercise 6 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp: skills: key:9007c6f402
## Exercise 7 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:MultipleChoiceExercise lang:r xp: skills: key:e6e6ab61f6
## Exercise 8 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```

--- type:VideoExercise lang:r xp: skills: key:a6b36fcc4b
## Exercise 9 
Content

*** =video_link

--- type:NormalExercise lang:r xp: skills: key:9e3e98db06
## Exercise 10 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp: skills: key:d13a898e6d
## Exercise 11 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
