---
title        : Chapter 2
description  : Description 

--- type:VideoExercise lang:r xp: skills: key:fbacaac39c
## Exercise 1 
Content

*** =video_link

--- type:NormalExercise lang:r xp: skills: key:c3a44c5f21
## Exercise 2 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp: skills: key:6f5ba1166d
## Exercise 3 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:MultipleChoiceExercise lang:r xp: skills: key:833bfe1555
## Exercise 4 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```

--- type:VideoExercise lang:r xp: skills: key:3ff33cc12c
## Exercise 5 
Content

*** =video_link

--- type:NormalExercise lang:r xp: skills: key:c3bb7de63e
## Exercise 6 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp: skills: key:b84231fbb3
## Exercise 7 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:MultipleChoiceExercise lang:r xp: skills: key:2f308dd147
## Exercise 8 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```

--- type:VideoExercise lang:r xp: skills: key:cc9a359ff1
## Exercise 9 
Content

*** =video_link

--- type:NormalExercise lang:r xp: skills: key:98077336b7
## Exercise 10 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp: skills: key:b5801555c8
## Exercise 11 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:MultipleChoiceExercise lang:r xp: skills: key:2bf77edf5d
## Exercise 12 
Content

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```
